# 1. Neural Activation
## Activation Function
With the advent of feedforword neural networks, activation functions have become an extremely important feature of the artificial neural networks. Activation functions are used to introduce non-lineraity in ANNs, and hence, determine whether a neuron is activated or not.

Most neural networks describe the features by using an affine transformation controlled by learned parameters, followed by an activation function.

A single layer in a neural network can be mathematically represented as:
$$H = \sigma (W*X + b)$$
where $W$ is a weight matrix, $X$ is the input and $b$ is the bias matrix. $*$ denotes the matrix multiplication and $\sigma (Y)$ is the activation function.

**Note**:  $\sigma (Y)$ is applied to every element of the matrix, Y.

There are many activation functions that exist, but for this problem we will implement two activation functions.

```python
#Import libraries
import numpy as np
```

### Part 1: Implement the affine transformation, $W*X + b$

<!-- grade_id: transformation, points: 0, grade: false, solution: true -->
```python
def transformation(W,X,b):
    """
    Implement the transformation W*X + b, given the matrices W, X, and b.
    
    Note that all matrix calculations follow the general matrix arithmatic rules.
    
    Parameters: W,X,b
    Output: transformed_X, i.e., W*X + b
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return transformed_X
```

<!-- grade_id: transformation-visible, points: 0, grade: true, solution: false -->
```python
print("Running base test case 1...")

W_test1 = np.array([[-4., 3., -6., 4.], 
                    [4., 0., -7., 6.], 
                    [-3., 6., -6., 0.]])

X_test1 = np.array([[2., 2.], 
                    [5., 3.], 
                    [0., 5.], 
                    [0., 5.]])

b_test1 = np.array([[0., 1.], 
                    [1., 0.], 
                    [1., 0.]])

test1 = transformation(W_test1, X_test1, b_test1)

ans_test1 = np.array([[  7.,  -8.],
                      [  9.,   3.],
                      [ 25., -18.]])

assert np.allclose(test1, ans_test1, rtol=1e-05, atol=1e-06)

print("Base test case 1 successful!!\n")



print("Running base test case 2...")

W_test2 = np.array([[ -0.7787005 ,  -0.47647797,  0.11260233],
                    [ -0.14420051,  0.17060967,  -0.6843165 ]])

X_test2 = np.array([[ 0.11699419,  0.42106442],
                    [ 0.9917111 ,  0.77009803],
                    [ 0.84847815,  0.51806326]])

b_test2 = np.array([[ 0.28954369,  0.33627522],
                    [ 0.5604489 ,  0.67298448]])

test2 = transformation(W_test2, X_test2, b_test2)

ans_test2 = np.array([[-0.17854762, -0.30020747],
                      [ 0.13214618,  0.38913371]])

assert np.allclose(test2, ans_test2, rtol=1e-05, atol=1e-06)

print("Base test case 2 successful!!\n")
```

<!-- grade_id: transformation-hidden, points: 2, grade: true, solution: false -->
```python
# Running hidden test case for transformation. Don't edit the cell.                            *** 2 marks ***
```

### Part 2: Implement the $tanh$ activation function

$$\sigma (x) = tanh(x) = \frac{2}{1 + e^{-2x}} - 1$$

<!-- grade_id: activation_tanh, points: 0, grade: false, solution: true -->
```python
def activation_tanh(Y):
    """
    Given a matrix Y, apply the tanh activation function to each element.
    
    Paramaters: Y
    Output: H
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
         
    return H
```

<!-- grade_id: tanh-visible, points: 0, grade: true, solution: false -->
```python
print("Running base test case 1...")

H_test1 = activation_tanh(ans_test1)

H_ans_test1 = np.array([[ 0.99999834, -0.99999977],
                        [ 0.99999997,  0.99505475],
                        [ 1.        , -1.        ]])

assert np.allclose(H_test1, H_ans_test1, rtol=1e-05, atol=1e-06)

print("Base test case 1 successful!!\n")



print("Running base test case 2...")

H_test2 = activation_tanh(ans_test2)

H_ans_test2 = np.array([[-0.17667418, -0.29150246],
                        [ 0.13138231,  0.37061317]])

assert np.allclose(H_test2, H_ans_test2, rtol=1e-05, atol=1e-06)

print("Base test case 2 successful!!\n")
```

<!-- grade_id: tanh-hidden, points: 2, grade: true, solution: false -->
```python
# # Running hidden test case for tanh. Don't edit the cell.                                     *** 2 marks ***
```

### Part 3: Implement the Exponential Linear Unit (ELU) activation function

$$ \sigma (x) = f(\alpha, x) = \begin{cases} \alpha(e^x -1) &\mbox{if } x < 0 \\ 
x & \mbox{if } x \geq 0 \end{cases} $$

<!-- grade_id: activation_elu, points: 0, grade: false, solution: true -->
```python
def activation_elu(Y, alpha):
    """
    Given a matrix, Y, and a real number, alpha, apply the ELU activation function to each element.
    
    Paramaters: Y, alpha
    Output: Z
    """
    
    # Hint: Use A = np.copy(B) to create deep copies of numpy array. A = B creates shallow copies of B.
    
    # YOUR CODE HERE
    raise NotImplementedError()
         
    return Z
```

<!-- grade_id: elu-visible, points: 0, grade: true, solution: false -->
```python
print("Running base test case 1...")

Z_test1 = activation_elu(ans_test1, 0.8)

Z_ans_test1 = np.array([[  7.        ,  -0.79973163],
                        [  9.        ,   3.        ],
                        [ 25.        ,  -0.79999999]])

assert np.allclose(Z_test1, Z_ans_test1, rtol=1e-05, atol=1e-06)

print("Base test case 1 successful!!\n")



print("Running base test case 2...")

Z_test2 = activation_elu(ans_test2, 1.)

Z_ans_test2 = np.array([[-0.16351578, -0.25933546],
                     [ 0.13214618,  0.38913371]])

assert np.allclose(Z_test2, Z_ans_test2, rtol=1e-05, atol=1e-06)

print("Base test case 2 successful!!\n")
```

<!-- grade_id: elu-hidden, points: 2, grade: true, solution: false -->
```python
# # Running hidden test case for ELU. Don't edit the cell.                                      *** 2 marks ***
```