```python
# Import libraries 
import numpy as np
```

 # The Convolution Conundrum

Convolutions are the fundamental operations that are performed in a Convolutional Neural Network. This question aims to realise these operations using our knowledge of numpy.

Refer to the figure below which illustrates how a convolution is performed. The explanation goes as follows:-

1. The kernel window (the yellow matrix) is superimposed over the 3x3 region in the source matrix and a dot product is computed between the values in the kernel and corresponding values in the source matrix.

NOTE: Dot product is computed between 2 matrices A and B as follows:-
      Multiply each element in A with the corresponding element in B.
      The Dot product is the sum of all these products.

2. Now the kernel is shifted to the right by one position and the operation described above is performed again to obtain another dot product value. 

3. This right shift happens continuously so long as the kernel does not cross the matrix bounds.

4. Once the kernel reaches the right most position possible for a given row it jumps to the leftmost position possible in the next row and starts moving right again, computing dot products along the way. Note that at no point does the kernel go outside the source matrix bounds. 

Before Convolution, it is common to perform Zero Padding on the source matrix as an additional operation in order to preserve the size of the source matrix after the convolution. Zero padding involves attaching rows and columns of zeros on all four sides of the source matrix. The number of rows and columns to attach depends on the kernel size and the desired output size and it is easy to think why. 

# Task

The task is divied into two sections:-
1. Perform appropriate Zero Padding on the imput image so as to obtain output of same size as original input matrix.
2. Perform the convolution operation on the matrix obtained after padding. The kernel is provided as a function parameter. Assume kernel dimension is an odd number.



# Marking 

You will be marked on the hidden test cases.  
Marks will be awarded for Zero Padding and Convolution seperately.  
4 marks for Zero Padding  
6 marks for Convolution


# CONVOLUTION

![](https://i.stack.imgur.com/I7DBr.gif)

# ZERO PADDING EXAMPLE

![](http://deeplearning.net/software/theano/_images/arbitrary_padding_no_strides.gif) 

<!-- grade_id: cell-d188d2d47db02790, points: 0, grade: false, solution: true -->
```python
def zero_pad(img, kernel):
    
    """ Both img and kernel are square matrices.
        Define a matrix zero_padded_img with appropriate size.
        zero_padded_img is returned by the function.
        Assume kernel dimension is an odd number. """
    
    ### BEGIN SOLUTION
    
    n = img.shape[0]
    k = kernel.shape[0]
    z = int(k/2)
    zero_padded_img = np.zeros(shape=(n+2*z,n+2*z))
    zero_padded_img[z:n+z,z:n+z] = img
    
    ### END SOLUTION
    
    return zero_padded_img
```

<!-- grade_id: cell-50fdca5027f11d90, points: 0, grade: false, solution: true -->
```python
def convolve_image(img, kernel):
    """ Both img and kernel are square matrices.
        Assume kernel dimension is an odd number. """
    
    ans_img = np.zeros_like(img)
    
    ### BEGIN SOLUTION
    n = img.shape[0]
    k = kernel.shape[0]
    zImg = zero_pad(img,kernel)
    
    for i in range(n):
        for j in range(n):
            ans_img[i,j] = np.sum( zImg[i:i+k,j:j+k] * kernel )
    ### END SOLUTION
    
    return ans_img
```

<!-- grade_id: cell-f7b89b5d1f21381b, points: 2, grade: true, solution: false -->
```python
# Running Hidden Test Case...
### BEGIN HIDDEN TESTS 
print("Zero-Padding: Running hidden test case 1...")
img_hidden_test1 = np.array([[ 0,  1, -4,  2,  4,  0,  4, -2,  4, -5,  3, -4, -3, -5,  1, -1, -3,
        -4,  1,  4],
       [-4,  4,  4, -3,  3,  3, -2,  2, -4,  2,  3,  3, -4,  4, -4, -2, -2,
        -5,  2, -1],
       [-2,  4,  2,  0,  1, -1, -3,  4, -4,  3, -4, -2,  1,  0,  2, -4,  4,
         1, -1,  2],
       [ 4,  0,  4,  3,  4,  3, -4, -1, -2, -5, -1,  0,  4,  2,  3,  0,  1,
         0, -1,  0],
       [-5,  0, -1,  2,  2, -1, -2, -5,  0,  2, -5, -1,  4,  3, -1,  0,  4,
        -2,  0,  4],
       [ 0,  3, -3, -3, -1, -5,  1,  0, -4, -1, -5,  2,  0,  3,  4, -2, -3,
        -5, -2,  3],
       [-2,  4, -4, -2, -4, -2, -3,  1,  4,  4, -5,  2,  1,  4,  0,  1,  1,
         4,  2, -2],
       [-5, -3,  2,  3, -3,  1,  4,  3,  3,  1, -1, -5,  0, -3, -1,  4, -3,
        -2, -2,  3],
       [ 2, -2, -1, -5, -3, -2,  0, -3, -3, -4, -4,  3,  3,  0,  3,  2, -3,
         0,  0, -2],
       [ 0, -5,  1, -4, -4,  4,  4,  3, -5,  2, -5, -3, -4,  0,  3, -1, -5,
        -1,  0, -5],
       [ 3,  0,  0,  3, -5, -1, -2,  3, -5,  3,  3,  4,  3,  1, -4,  3, -4,
        -1,  3, -2],
       [-1, -2,  2,  0,  1, -4,  2, -4, -3, -5,  1,  2, -5,  3, -4,  4, -4,
        -1, -3, -2],
       [ 2,  0, -5,  2, -4,  0,  0, -2,  4, -1, -3,  0,  3, -3, -4,  4, -4,
        -5,  3, -5],
       [ 4, -1,  1, -1, -1,  3,  4, -3,  2,  0, -3,  1,  3, -1,  0,  4,  3,
         3, -3,  4],
       [-5, -2, -1,  4, -3,  2,  2,  2,  2,  2, -5, -1,  0,  3,  1,  0,  2,
        -2,  0, -4],
       [-1, -2,  0, -5, -2, -1, -2,  3, -1, -3, -4,  0, -2, -2, -5, -1, -1,
        -5, -5, -5],
       [-3, -1, -2, -5,  3,  2,  2, -4,  2, -1, -3,  3, -2,  0, -5, -4,  1,
         2,  1,  3],
       [-5, -1, -4, -1, -3, -2,  2,  2, -2, -5,  4, -2, -3,  3,  3,  0, -2,
         1,  4,  4],
       [-5, -3, -4,  3, -3, -5,  2,  1,  3, -2,  2,  4,  2, -3,  3,  0, -2,
        -4,  3, -3],
       [-4,  1,  3, -4, -3, -1,  4, -1, -2, -4, -5,  4,  0,  2,  4, -5, -1,
         2, -5, -5]])
ker_hidden_test1 = np.array([[-2 , -4,  4,  2 , 4],
 [-3 ,-3 ,-3 ,-5  ,0],
 [-5 ,-4 ,-5 ,-5  ,4],
 [ 3  ,3  ,1 ,-5  ,3],
 [ 0 ,-3  ,4  ,0 ,-1]])
zp_hidden_test1 = zero_pad(img_hidden_test1,ker_hidden_test1)
ans_zp_hidden_test1 = np.array([[ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  1., -4.,  2.,  4.,  0.,  4., -2.,  4., -5.,  3.,
        -4., -3., -5.,  1., -1., -3., -4.,  1.,  4.,  0.,  0.],
       [ 0.,  0., -4.,  4.,  4., -3.,  3.,  3., -2.,  2., -4.,  2.,  3.,
         3., -4.,  4., -4., -2., -2., -5.,  2., -1.,  0.,  0.],
       [ 0.,  0., -2.,  4.,  2.,  0.,  1., -1., -3.,  4., -4.,  3., -4.,
        -2.,  1.,  0.,  2., -4.,  4.,  1., -1.,  2.,  0.,  0.],
       [ 0.,  0.,  4.,  0.,  4.,  3.,  4.,  3., -4., -1., -2., -5., -1.,
         0.,  4.,  2.,  3.,  0.,  1.,  0., -1.,  0.,  0.,  0.],
       [ 0.,  0., -5.,  0., -1.,  2.,  2., -1., -2., -5.,  0.,  2., -5.,
        -1.,  4.,  3., -1.,  0.,  4., -2.,  0.,  4.,  0.,  0.],
       [ 0.,  0.,  0.,  3., -3., -3., -1., -5.,  1.,  0., -4., -1., -5.,
         2.,  0.,  3.,  4., -2., -3., -5., -2.,  3.,  0.,  0.],
       [ 0.,  0., -2.,  4., -4., -2., -4., -2., -3.,  1.,  4.,  4., -5.,
         2.,  1.,  4.,  0.,  1.,  1.,  4.,  2., -2.,  0.,  0.],
       [ 0.,  0., -5., -3.,  2.,  3., -3.,  1.,  4.,  3.,  3.,  1., -1.,
        -5.,  0., -3., -1.,  4., -3., -2., -2.,  3.,  0.,  0.],
       [ 0.,  0.,  2., -2., -1., -5., -3., -2.,  0., -3., -3., -4., -4.,
         3.,  3.,  0.,  3.,  2., -3.,  0.,  0., -2.,  0.,  0.],
       [ 0.,  0.,  0., -5.,  1., -4., -4.,  4.,  4.,  3., -5.,  2., -5.,
        -3., -4.,  0.,  3., -1., -5., -1.,  0., -5.,  0.,  0.],
       [ 0.,  0.,  3.,  0.,  0.,  3., -5., -1., -2.,  3., -5.,  3.,  3.,
         4.,  3.,  1., -4.,  3., -4., -1.,  3., -2.,  0.,  0.],
       [ 0.,  0., -1., -2.,  2.,  0.,  1., -4.,  2., -4., -3., -5.,  1.,
         2., -5.,  3., -4.,  4., -4., -1., -3., -2.,  0.,  0.],
       [ 0.,  0.,  2.,  0., -5.,  2., -4.,  0.,  0., -2.,  4., -1., -3.,
         0.,  3., -3., -4.,  4., -4., -5.,  3., -5.,  0.,  0.],
       [ 0.,  0.,  4., -1.,  1., -1., -1.,  3.,  4., -3.,  2.,  0., -3.,
         1.,  3., -1.,  0.,  4.,  3.,  3., -3.,  4.,  0.,  0.],
       [ 0.,  0., -5., -2., -1.,  4., -3.,  2.,  2.,  2.,  2.,  2., -5.,
        -1.,  0.,  3.,  1.,  0.,  2., -2.,  0., -4.,  0.,  0.],
       [ 0.,  0., -1., -2.,  0., -5., -2., -1., -2.,  3., -1., -3., -4.,
         0., -2., -2., -5., -1., -1., -5., -5., -5.,  0.,  0.],
       [ 0.,  0., -3., -1., -2., -5.,  3.,  2.,  2., -4.,  2., -1., -3.,
         3., -2.,  0., -5., -4.,  1.,  2.,  1.,  3.,  0.,  0.],
       [ 0.,  0., -5., -1., -4., -1., -3., -2.,  2.,  2., -2., -5.,  4.,
        -2., -3.,  3.,  3.,  0., -2.,  1.,  4.,  4.,  0.,  0.],
       [ 0.,  0., -5., -3., -4.,  3., -3., -5.,  2.,  1.,  3., -2.,  2.,
         4.,  2., -3.,  3.,  0., -2., -4.,  3., -3.,  0.,  0.],
       [ 0.,  0., -4.,  1.,  3., -4., -3., -1.,  4., -1., -2., -4., -5.,
         4.,  0.,  2.,  4., -5., -1.,  2., -5., -5.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.]])
assert np.array_equal(ans_zp_hidden_test1,zp_hidden_test1)
print("Zero-Padding: Hidden test case 1 successful!!")
### END HIDDEN TESTS 
```

<!-- grade_id: cell-8ab150b4841868d2, points: 2, grade: true, solution: false -->
```python
# Running Hidden Test Case...
### BEGIN HIDDEN TESTS 
print("Zero_Padding: Running hidden test case 2...")
img_hidden_test2 = np.array([[-3, -5, -1, -1,  4,  2,  0, -5,  2,  0, -5],
       [-4, -2, -1,  4,  2,  1,  1,  1,  0,  3, -2],
       [-3,  1,  1, -4, -1, -4, -5, -3,  0, -4, -5],
       [ 0, -1,  1,  3, -5, -3,  0, -3,  2, -2, -4],
       [ 4, -3,  2, -1,  0,  3,  4,  2,  2, -5, -4],
       [ 4,  1, -1,  4,  4, -1,  3,  4,  4,  1,  4],
       [ 3,  1, -1,  0, -2,  4,  4,  3, -3,  3,  3],
       [-2, -2, -4, -4, -4, -5, -5,  3, -2,  2,  2],
       [-2, -4,  3,  1,  1, -5, -2,  2, -3, -3, -5],
       [-2,  3, -3,  1,  4, -1, -1,  1, -1, -2,  2],
       [ 2,  4, -3, -3,  4, -3, -1, -2, -5,  0,  2]])
ker_hidden_test2 = np.array([[-3,  3, -3, -3,  2, -4, -1],
 [ 0, -5,  0, -2,  2,  4, -4],
 [-1, -5,  3, -1, -1, -1,  3],
 [-4, -1,  1, -2,  4, -5,  0],
 [-5,  1,  0, -1,  4, -2,  2],
 [-5, -1, -4, -3,  3,  0,  1],
 [-5,  1, -4, -3, -2, -2,  2]])
zp_hidden_test2 = zero_pad(img_hidden_test2,ker_hidden_test2)
ans_zp_hidden_test2 = np.array([[ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.],
       [ 0.,  0.,  0., -3., -5., -1., -1.,  4.,  2.,  0., -5.,  2.,  0.,
        -5.,  0.,  0.,  0.],
       [ 0.,  0.,  0., -4., -2., -1.,  4.,  2.,  1.,  1.,  1.,  0.,  3.,
        -2.,  0.,  0.,  0.],
       [ 0.,  0.,  0., -3.,  1.,  1., -4., -1., -4., -5., -3.,  0., -4.,
        -5.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0., -1.,  1.,  3., -5., -3.,  0., -3.,  2., -2.,
        -4.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  4., -3.,  2., -1.,  0.,  3.,  4.,  2.,  2., -5.,
        -4.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  4.,  1., -1.,  4.,  4., -1.,  3.,  4.,  4.,  1.,
         4.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  3.,  1., -1.,  0., -2.,  4.,  4.,  3., -3.,  3.,
         3.,  0.,  0.,  0.],
       [ 0.,  0.,  0., -2., -2., -4., -4., -4., -5., -5.,  3., -2.,  2.,
         2.,  0.,  0.,  0.],
       [ 0.,  0.,  0., -2., -4.,  3.,  1.,  1., -5., -2.,  2., -3., -3.,
        -5.,  0.,  0.,  0.],
       [ 0.,  0.,  0., -2.,  3., -3.,  1.,  4., -1., -1.,  1., -1., -2.,
         2.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  2.,  4., -3., -3.,  4., -3., -1., -2., -5.,  0.,
         2.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.],
       [ 0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,  0.,
         0.,  0.,  0.,  0.]])
assert np.array_equal(ans_zp_hidden_test2,zp_hidden_test2)
print("Zero-Padding: Hidden test case 2 successful!!")
### END HIDDEN TESTS 
```

<!-- grade_id: cell-5925c1058171e6c7, points: 3, grade: true, solution: false -->
```python
# Running Hidden Test Case...
### BEGIN HIDDEN TESTS 
print("Convolution: Running hidden test case 1...")
img_hidden_test1 = np.array([[ 0,  1, -4,  2,  4,  0,  4, -2,  4, -5,  3, -4, -3, -5,  1, -1, -3,
        -4,  1,  4],
       [-4,  4,  4, -3,  3,  3, -2,  2, -4,  2,  3,  3, -4,  4, -4, -2, -2,
        -5,  2, -1],
       [-2,  4,  2,  0,  1, -1, -3,  4, -4,  3, -4, -2,  1,  0,  2, -4,  4,
         1, -1,  2],
       [ 4,  0,  4,  3,  4,  3, -4, -1, -2, -5, -1,  0,  4,  2,  3,  0,  1,
         0, -1,  0],
       [-5,  0, -1,  2,  2, -1, -2, -5,  0,  2, -5, -1,  4,  3, -1,  0,  4,
        -2,  0,  4],
       [ 0,  3, -3, -3, -1, -5,  1,  0, -4, -1, -5,  2,  0,  3,  4, -2, -3,
        -5, -2,  3],
       [-2,  4, -4, -2, -4, -2, -3,  1,  4,  4, -5,  2,  1,  4,  0,  1,  1,
         4,  2, -2],
       [-5, -3,  2,  3, -3,  1,  4,  3,  3,  1, -1, -5,  0, -3, -1,  4, -3,
        -2, -2,  3],
       [ 2, -2, -1, -5, -3, -2,  0, -3, -3, -4, -4,  3,  3,  0,  3,  2, -3,
         0,  0, -2],
       [ 0, -5,  1, -4, -4,  4,  4,  3, -5,  2, -5, -3, -4,  0,  3, -1, -5,
        -1,  0, -5],
       [ 3,  0,  0,  3, -5, -1, -2,  3, -5,  3,  3,  4,  3,  1, -4,  3, -4,
        -1,  3, -2],
       [-1, -2,  2,  0,  1, -4,  2, -4, -3, -5,  1,  2, -5,  3, -4,  4, -4,
        -1, -3, -2],
       [ 2,  0, -5,  2, -4,  0,  0, -2,  4, -1, -3,  0,  3, -3, -4,  4, -4,
        -5,  3, -5],
       [ 4, -1,  1, -1, -1,  3,  4, -3,  2,  0, -3,  1,  3, -1,  0,  4,  3,
         3, -3,  4],
       [-5, -2, -1,  4, -3,  2,  2,  2,  2,  2, -5, -1,  0,  3,  1,  0,  2,
        -2,  0, -4],
       [-1, -2,  0, -5, -2, -1, -2,  3, -1, -3, -4,  0, -2, -2, -5, -1, -1,
        -5, -5, -5],
       [-3, -1, -2, -5,  3,  2,  2, -4,  2, -1, -3,  3, -2,  0, -5, -4,  1,
         2,  1,  3],
       [-5, -1, -4, -1, -3, -2,  2,  2, -2, -5,  4, -2, -3,  3,  3,  0, -2,
         1,  4,  4],
       [-5, -3, -4,  3, -3, -5,  2,  1,  3, -2,  2,  4,  2, -3,  3,  0, -2,
        -4,  3, -3],
       [-4,  1,  3, -4, -3, -1,  4, -1, -2, -4, -5,  4,  0,  2,  4, -5, -1,
         2, -5, -5]])
ker_hidden_test1 = np.array([[-2 , -4,  4,  2 , 4],
 [-3 ,-3 ,-3 ,-5  ,0],
 [-5 ,-4 ,-5 ,-5  ,4],
 [ 3  ,3  ,1 ,-5  ,3],
 [ 0 ,-3  ,4  ,0 ,-1]])
hidden_test1 = convolve_image(img_hidden_test1,ker_hidden_test1)
ans_hidden_test1 = np.array([[ -43.,    8.,   45.,   -9.,    0.,  -46.,  -25.,    7.,  -24.,
           4.,  -63.,   82.,   35.,   64.,   27.,   -5.,   78.,    8.,
         -15.,   -3.],
       [   7.,  -46.,   33.,  -34.,  -38.,   -2., -115.,   19.,   -1.,
         -10.,   -7.,   23.,   -4.,   -5.,   94.,    2.,   95.,   57.,
          23.,   24.],
       [ -27.,  -24.,   -2.,  -21.,  -15.,   49.,  -13.,  -21.,   33.,
         -64.,  -12.,  -66.,   17.,  -12.,   66.,    7.,   35.,   29.,
          33.,   34.],
       [ -15.,   -1.,  -77.,  -92.,  -65.,  -16.,  -15.,   16.,  -12.,
         107.,   79.,   24.,  -82.,   -7.,  -22., -100.,  -22.,    3.,
           1.,   12.],
       [ -11.,   62.,  -17.,  -86.,  -76.,  -50.,  -32.,  110.,   17.,
          89.,  -21.,   21.,  -15.,  -71.,  -23.,  -52.,   30.,   28.,
         -89.,  -29.],
       [ -36.,   24.,   72.,   17.,   -3.,   26.,  -18.,   20.,  -32.,
         126.,   94.,   41.,   31.,  -68.,  -31.,  -25.,  -17.,   49.,
          23.,   50.],
       [ -40.,    4.,    0.,   47.,  106.,   31.,   52.,   41.,  -10.,
          35.,   18.,   55.,   12.,  -81.,  -65.,  -16.,   13.,   16.,
         -30.,   -3.],
       [  36.,   23.,   20.,  -19.,   33.,    5.,   20.,  -69., -150.,
          -6.,  -65.,   13.,   45.,   70.,  -37.,  -56.,  -55.,    7.,
          35.,   -1.],
       [  50.,  -25.,  -30.,   68.,  -25.,   -1.,    1.,  102.,  -39.,
          94.,   17.,   73.,   37.,  -40., -106.,   77.,   -4.,  -11.,
          39.,  -36.],
       [  12.,   46.,   63.,  147.,   59.,   65.,  -12.,   27.,  -18.,
           3.,   45.,   27.,   60.,   65.,  -74.,   66.,  -43.,    2.,
          74.,   65.],
       [  38.,  -51.,  -60.,   30.,   28.,   17.,  -27.,  -20.,   24.,
         -19.,   -7.,   58.,  -18.,   -2.,  -69.,   72.,   -6.,    9.,
          83.,  -30.],
       [  10.,  -17.,  -30.,   -8.,   28.,   75.,   27.,  -11.,    5.,
          55.,  -63.,  -23.,  -30.,   72.,  -56.,  -12.,   46.,  -39.,
          49.,   23.],
       [ -12.,   30.,  -11.,   53.,  -24.,   -2.,   92.,   51.,   14.,
         111.,   33.,   11.,  -30.,   82.,  -69.,   30.,   67.,   49.,
         105.,   26.],
       [ -19.,   -1.,  -42.,   31.,   37.,  -64.,   11.,  -51.,  -57.,
          75.,   33.,   14.,  -57.,   54.,  -50.,   37.,  -40.,  -46.,
          22.,  -11.],
       [  11.,   17.,  -35.,   25.,  -17.,  -39.,  -17.,  -87.,  -29.,
          -1.,  -23.,   30.,  -15.,  -16.,  -89.,  -86.,  -57.,  -54.,
          19.,  -30.],
       [  38.,   -5.,   38.,   44.,    4.,   28.,   59.,  -65.,  -52.,
          32.,   48.,   41.,    9.,   63.,   67.,    6.,  -15.,   44.,
          81.,  126.],
       [ -31.,   43.,   80.,  110.,    0.,    5.,    4.,   -2.,   38.,
          -3.,  -28.,   99.,   35.,   10.,  111.,   84.,   81.,    7.,
          26.,    2.],
       [  -1.,   70.,   37.,  -24.,   82.,   31.,   -1.,  -61.,   32.,
          -5.,  -28.,   83.,   10.,   13.,   12.,  -29.,   52.,  -65.,
         -95.,  -60.],
       [  22.,   51.,   74.,   68.,  149.,  -27.,   16.,  -29.,   26.,
          20.,  -67.,  -11.,  -56.,  -82.,  -18.,   27.,   28.,    4.,
          20.,  -14.],
       [  19.,   28.,   14.,   41.,   78.,   36.,   27.,  -41.,  -28.,
          42.,   26.,   -2.,   12.,  -60.,  -15.,    4.,   -4.,   63.,
          95.,   45.]])
assert np.array_equal(ans_hidden_test1,hidden_test1)
print("Convolution: Hidden test case 1 successful!!")
### END HIDDEN TESTS 
```

<!-- grade_id: cell-9de25b6390f5a92a, points: 3, grade: true, solution: false -->
```python
# Running Hidden Test Case...
### BEGIN HIDDEN TESTS 
print("Convolution: Running hidden test case 2...")
img_hidden_test2 = np.array([[-3, -5, -1, -1,  4,  2,  0, -5,  2,  0, -5],
       [-4, -2, -1,  4,  2,  1,  1,  1,  0,  3, -2],
       [-3,  1,  1, -4, -1, -4, -5, -3,  0, -4, -5],
       [ 0, -1,  1,  3, -5, -3,  0, -3,  2, -2, -4],
       [ 4, -3,  2, -1,  0,  3,  4,  2,  2, -5, -4],
       [ 4,  1, -1,  4,  4, -1,  3,  4,  4,  1,  4],
       [ 3,  1, -1,  0, -2,  4,  4,  3, -3,  3,  3],
       [-2, -2, -4, -4, -4, -5, -5,  3, -2,  2,  2],
       [-2, -4,  3,  1,  1, -5, -2,  2, -3, -3, -5],
       [-2,  3, -3,  1,  4, -1, -1,  1, -1, -2,  2],
       [ 2,  4, -3, -3,  4, -3, -1, -2, -5,  0,  2]])
ker_hidden_test2 = np.array([[-3,  3, -3, -3,  2, -4, -1],
 [ 0, -5,  0, -2,  2,  4, -4],
 [-1, -5,  3, -1, -1, -1,  3],
 [-4, -1,  1, -2,  4, -5,  0],
 [-5,  1,  0, -1,  4, -2,  2],
 [-5, -1, -4, -3,  3,  0,  1],
 [-5,  1, -4, -3, -2, -2,  2]])

hidden_test2 = convolve_image(img_hidden_test2,ker_hidden_test2)
ans_hidden_test2 = np.array([[  11.,   -2.,  -34.,   62.,   38.,   77.,   -8.,   36.,   78.,
           0.,  108.],
       [  -4.,  -13.,   13.,   -7.,  -15.,    2.,  -67.,  -31.,   21.,
           9.,   80.],
       [  -6.,  -29.,   28.,   23.,   32.,  -44., -108.,   -2.,  -66.,
         -63.,   37.],
       [ -68.,   -9.,   91.,  -48.,  -14.,    5., -115.,   55.,  -17.,
         -99.,  -31.],
       [  29.,  -12.,   22.,   21.,   33.,   63.,   47.,   54.,  -24.,
          24.,  -96.],
       [   5.,   56.,  -33.,   35.,   78.,    5.,  131.,   71.,   28.,
          87.,  -50.],
       [  -3.,   48.,  -52.,  -42.,    7.,   19.,   53.,  -18.,   -3.,
           3.,   -9.],
       [ -54.,   20.,   10.,   14.,  -14.,  -95.,  104.,  -42.,   60.,
           0.,   -6.],
       [ -24.,  -62.,  -66.,   68.,  -42.,   21.,   27.,  -91.,   90.,
          -1.,    8.],
       [  48.,  -46.,  -21.,   49.,  -77.,   40.,    2.,  -51.,   72.,
           4.,  -27.],
       [  58.,   40.,   45.,   42.,  -33.,   67.,   40.,  -12.,    4.,
          16.,    3.]])
assert np.array_equal(ans_hidden_test2,hidden_test2)
print("Convolution: Hidden test case 2 successful!!")
### END HIDDEN TESTS 
```

```python
print("Running base test case 1...")
img_test1 = np.array([[1, 1, 1, 0, 0],
                      [0, 1, 1, 1, 0],
                      [0, 0, 1, 1, 1],
                      [0, 0, 1, 1, 0],
                      [0, 1, 1, 0, 0]])
ker_test1 = np.array([[1, 0, 1],
                      [0, 1, 0],
                      [1, 0, 1]])
test1 = convolve_image(img_test1,ker_test1)
ans_test1 = np.array([[2, 2, 3, 1, 1],
                      [1, 4, 3, 4, 1],
                      [1, 2, 4, 3, 3],
                      [1, 2, 3, 4, 1],
                      [0, 2, 2, 1, 1]])
assert np.array_equal(ans_test1,test1)
print("Base test case 1 successful!!")
```

```python
print("Running base test case 2...")
img_test2 = np.array([[-1, 0, 2, 2, 3, 1],
                     [2, 1, 4, 3, 4, 1],
                     [-3, 1, 2, 4, 3, 3],
                     [0, 1, 2, 3, 4, 1],
                     [4, 0, 2, 2, 1, 1],
                     [5, 1, 3, 2, 2, 0]])
ker_test2 = np.array([[1, 0, 1],
                      [2, 4, 1],
                      [2, 1, 0]])
test2 = convolve_image(img_test2,ker_test2)
ans_test2 = np.array([[ -2, 5,  16,  26,  27,  19],
                      [  6, 8,  27,  37,  37,  24],
                      [-10, 7,  22,  38,  37,  31],
                      [  6, 13,  20,  31,  35,  18],
                      [ 22, 23,  19,  27,  19,  14],
                      [ 21, 23,  18,  19,  15,   5]])
assert np.array_equal(ans_test2,test2)
print("Base test case 2 successful!!")
```

```python
print("Running base test case 3...")
img_test3 = np.array([[-1, 0, 2, 2, 3, 1, 1],
                [2, 1, 4, 3, 4, 1, 0],
                [-3, 1, 2, 4, 3, 3, 1],
                [0, 1, 2, 3, 4, 1, 4],
                [4, 0, 2, 2, 1, 1, 5],
                [5, 1, 3, 2, 2, 0, -1],
                [-2, 1, 0, -3, -3, 4, 0]])
ker_test3 = np.array([[1, 0, 1, 0, 0],
                      [2, 0, 1, 0, 1],
                      [2, 1, 0, 1, 1],
                      [0, 0, 2, 1, 2],
                      [1, 0, 3, 2, 0]])
test3 = convolve_image(img_test3,ker_test3)
ans_test3 = np.array([[  8, 22, 33, 37, 34, 25, 13],
                      [  7, 30, 41, 49, 47, 36, 34],
                      [ 25, 21, 44, 38, 53, 43, 45],
                      [ 33, 26, 39, 40, 50, 36, 29],
                      [ 14, 25, 20,  7, 30, 32, 14],
                      [  7,  9, 19, 13, 21, 22, 19],
                      [ 13, -2, 12,  9, 11, -2, 7]])
assert np.array_equal(ans_test3,test3)
print("Base test case 3 successful!")
```