## 2. Non-Negative Matrix Factorization

### Marks: 6


Non-negative matrix factorization is a popular unsupervised learning algorithm used for clustering, deconvolution and generating a lower-level representation of the data. In this problem you have to implement the classical multiplicative update rules proposed by Lee and Seung in 2001 (NIPS) to the solve the factorization problem.

### Problem:

Given a non-negative matrix $V \in m \times n $, the goal is to determine two non-negative factors of a lower rank $k$, i.e.  $W \in m \times k $ and $H \in k \times n$ such that:

$$ V \approx WH $$

To determine the matrices $W$ and $H$, one of the most common approaches is to minimize the frobenius norm between the original matrix and the reconstructed matrix, under the non-negativity constraint:

$$ \left \| V-WH \right \|^{2}  \ \ \ \ s.t. \ \ V,W,H \geq  0$$



### Task:

Your task in this question is only to implement the update rules for the above constrained problem. To minimize the frobenius norm under the given constraints, Lee and Seung proposed the following multiplicative update rules to update the matrices $W$ and $H$:


$$ W \leftarrow W \ast   \frac{(VH^T)}{(WHH^T)} $$

$$ H \leftarrow H \ast \frac{(W^TV)}{(W^TWH)} $$

where $\ast$ denotes element-wise multiplication. The division operation is element-wise and $AB$ denotes matrix multiplication between $A$ and $B$. 

To solve the non-negative matrix factorization problem we generally initialize matrices $W$ and $H$ randomly according to the required lower rank $k$ and then apply the update rules iteratively until convergence.


```python
# Import libraries 
import numpy as np
import numpy.linalg as LA

```

```python
# Base Function
# You do not have to modify this

def compute_NMF(V,W,H):
    """
    Given the matrices V and randomly initialized W and H, this function calls the functions 
    for updates rules iteratively and returns the updated matrices W, H and the frobenius error.
    """

    for i in range(50):
        W = update_W(V,W,H)
        H = update_H(V,W,H)

    error= LA.norm(V - np.dot(W,H))

    return W,H,error
```

### Part 1 : Implement update rules for W

<!-- grade_id: Update_W, points: 0, grade: false, solution: true -->
```python
def update_W(V,W,H):
    """
    Implement the update rules for W, given the matrices W, H and V
    Parameters : Matrices V, W and H
    Output: Matrix W with updated Values    
    """

    
    ### BEGIN SOLUTION
    VH = V.dot(H.T)
    WHtH = W.dot(H.dot(H.T))
    W *= VH
    W /= WHtH
    ### END SOLUTION
    
    return W
```

### Part 2 : Implement update rules for H

<!-- grade_id: Update_H, points: 0, grade: false, solution: true -->
```python
def update_H(V,W,H):
    """
    Implement the update rules for H, given the matrices W, H and V
    Parameters : Matrices V, W and H
    Output: Matrix H with updated Values
    
    """
    
    
    ### BEGIN SOLUTION
    VtW = np.dot(W.T,V)
    HWtW = np.dot(W.T.dot(W),H)
    H *= VtW
    H /= HWtW
    ### END SOLUTION
    
    return H
```

<!-- grade_id: case_1, points: 0, grade: true, solution: false -->
```python
print("Running base test case 1...\n")
V= np.array(  [[ 51.,  84.,  86.,  67.,  11.],
               [ 57.,  37.,  62.,  87.,   6.],
               [ 31.,  10.,  79.,  39.,  85.],
               [ 84.,  59.,  46.,  80.,  36.],
               [ 77.,  58.,  28.,  94.,  69.]])

W = np.array( [[  5.,   8.,  19.],
               [ 46.,   9.,  31.],
               [ 19.,   2.,  14.],
               [ 41.,  16.,   5.],
               [ 20.,  38.,   6.]])

H = np.array( [[ 34.,  34.,  43.,   3.,   6.],
               [ 33.,  18.,  15.,  29.,   9.],
               [  2.,  22.,   3.,   4.,   6.]])

ans_W = np.array( [[ 0.08544129,  0.64652928,  3.38656155],
                   [ 0.17424102,  0.9338013 ,  1.89953141],
                   [ 2.30985524,  0.0704793 ,  0.14674537],
                   [ 0.59893525,  1.35729956,  0.91280561],
                   [ 1.13200231,  1.5201163 ,  0.01445199]])

ans_H = np.array( [[  1.11310981e+01,   3.68721639e+00,   3.08251281e+01, 1.42221620e+01,   3.79139869e+01],
                   [  4.66804619e+01,   3.05525948e+01,   1.64087722e+00, 5.09355415e+01,   1.11386279e+01],
                   [  6.28131421e+00,   1.53952944e+01,   2.55431038e+01, 1.14672560e+01,   5.64956239e-09]])


error= LA.norm(V-np.dot(W,H))
print("Initial Frobenius Error: " + str(error))
W, H, error = compute_NMF(V,W,H)
print("Final Frobenius Error: " + str(error))

assert LA.norm(W - ans_W) < 10**-5
assert LA.norm(H - ans_H) < 10**-5

print("\nBase test case 1 successful!!")
```

Output:
```
Running base test case 1...

Initial Frobenius Error: 6067.71604477
Final Frobenius Error: 38.7679913414

Base test case 1 successful!!
```

<!-- grade_id: case_2, points: 0, grade: true, solution: false -->
```python
print("Running base test case 2...\n")

V= np.load('./inputs/V_2.npy')
W= np.load('./inputs/W_2.npy')
H= np.load('./inputs/H_2.npy')
ans_W= np.load('./inputs/W_2_ans.npy')
ans_H= np.load('./inputs/H_2_ans.npy')



error= LA.norm(V-np.dot(W,H))
print("Initial Frobenius Error: " + str(error))
W, H, error = compute_NMF(V,W,H)
print("Final Frobenius Error: " + str(error))

assert LA.norm(W - ans_W) < 10**-5
assert LA.norm(H - ans_H) < 10**-5

print("\nBase test case 2 successful!!")
```

Output:
```
Running base test case 2...

Initial Frobenius Error: 35195.0462423
Final Frobenius Error: 97.4290913594

Base test case 2 successful!!
```

<!-- grade_id: case_3, points: 6, grade: true, solution: false -->
```python
# RUNNING HIDDEN TEST CASE
### BEGIN HIDDEN TESTS
print("Running base test case 3...\n")

V= np.load('./inputs/V_3.npy')
W= np.load('./inputs/W_3.npy')
H= np.load('./inputs/H_3.npy')
ans_W= np.load('./inputs/W_3_ans.npy')
ans_H= np.load('./inputs/H_3_ans.npy')



error= LA.norm(V-np.dot(W,H))
print("Initial Frobenius Error: " + str(error))
W, H, error = compute_NMF(V,W,H)
print("Final Frobenius Error: " + str(error))

assert LA.norm(W - ans_W) < 10**-5
assert LA.norm(H - ans_H) < 10**-5
print("\nBase test case 3 successful!!")
### END HIDDEN TESTS
```

Output:
```
Running base test case 3...

Initial Frobenius Error: 253525.880028
Final Frobenius Error: 188.118488805

Base test case 3 successful!!
```