# 3. Nearest Redi Problem
### Marks : 8
In a hypothetical $n$-dimensional universe, there exists $p$ population of a particular species of human, Homo BITSians. These species like to hangout in specialized eateries, called Redi. In this universe, there are $q$ Redis which serve delicious snacks and beverages at nominal prices. Our task is to find the nearest Redi from each of the Homo BITSians so that they spend less time on commuting. Another part of the probelem is to give the number of Redis inside a radius of $r$ metres from each Homo BITSians which will help them to hangout in as many Redis as possible.
### Problem :
Matrices, $X \in p \times n$ and $Y \in q \times n$, which has the co-ordinates of $p$ Homo BITSians and $q$ Redis respectively in the $n$-dimensional universe are given. The $i^{th}$ row in the matrix, $X$, corresponds to the $i^{th}$ Homo BITSian. Similarly, the $i^{th}$ row in the matrix, $Y$, corresponds to the $i^{th}$ Redi.

**Note**: Here, row numbering (indexing) starts from $0$.
### Task:
In this question there are two tasks which need to completed :
1. Given $X$, $Y$, find a vector, $V$, of length $p$. The vector, $V$, is such that the $i^{th}$ element of $V$ has the index of the nearest Redi from the $i^{th}$ Homo BITSian.

2. Given $X$, $Y$ and $r$, find a vector, $W$, of length $p$. The vector, $W$, is such that the $i^{th}$ element of $W$ has the number of Redis inside a radius of $r$ (strictly less than), centered at the $i^{th}$ Homo BITSian.


Distance metric is the usual $l_2$-norm.
In a n-dimensional space with points $x = (x_0, x_0, \ldots, x_{n-1})$ and $y = (y_0, y_0, \ldots, y_{n-1})$, the distance can be calculated as:

$$D_{xy}^2 = (x_0 - y_0)^2 
+ (x_1 - y_1)^2 + \ldots + (x_{n-1} - y_{n-1})^2$$

```python
import numpy as np
```

### Part 1: Find the index of the nearest Redi from each Homo BITSian

<!-- grade_id: distances, points: 0, grade: false, solution: true -->
```python
# Base Distance Function to be completed by the student

def distances(X,Y):
    """
    Given matrices X and Y, the function returns a distance matrix. 
    The (i,j)th element of the matrix contains the distance of jth Redi 
    from the ith Homo BITSian.
    
    Parameters: X,Y
    Returns: D
    """
    
    ### BEGIN SOLUTION
    diff = X.reshape(X.shape[0], 1, X.shape[1]) - Y
    D_squared = np.sum(np.square(diff), axis=2)
    D = np.sqrt(D_squared)
    ### END SOLUTION
    
    return D
```

<!-- grade_id: nearest_redi, points: 0, grade: false, solution: true -->
```python
def nearest_redi(X,Y):
    """
    Given matrices X and Y, the function returns a nearest redi vector. 
    The i-th element of the vector contains the index of nearest Redi 
    from the ith Homo BITSian.
    
    Parameters: X,Y
    Returns: V
    """
    
    D = distances(X,Y)
    
    ### BEGIN SOLUTION
    V = np.argmin(D,1)
    ### END SOLUTION
    
    return V
```

<!-- grade_id: visible-task1, points: 0, grade: true, solution: false -->
```python
print("Running base test case 1...")

X_test1 = np.array([[-3.,  4.],
                    [ 4., -2.],
                    [-1.,  0.]])

Y_test1 = np.array([[-3.,  0.],
                    [-3., -3.]])

V_test1 = nearest_redi(X_test1, Y_test1)
V_ans_test1 = np.array([0, 1, 0])

assert np.array_equal(V_test1, V_ans_test1)

print("Base test case 1 successful!!\n")



print("Running base test case 2...")

X_test2 = np.array([[ 0.08170274, -4.8955951 , -4.0473417 ],
                    [-1.13259313,  4.38171415, -3.22068891]])

Y_test2 = np.array([[ 3.79010736,  1.70042849, -3.06603884],
                    [ 3.8921235 , -1.85207272,  2.33340715],
                    [ 1.67360485,  2.11437547,  0.87529999]])

V_test2 = nearest_redi(X_test2, Y_test2)
V_ans_test2 = np.array([0, 2])

assert np.array_equal(V_test2, V_ans_test2)

print("Base test case 2 successful!!\n")
```

Output:
```
Running base test case 1...
Base test case 1 successful!!

Running base test case 2...
Base test case 2 successful!!
```

<!-- grade_id: hidden-task1, points: 5, grade: true, solution: false -->
```python
# Running hidden test case for Part 1. Don't edit the cell.                                     *** 5 marks ***
### BEGIN HIDDEN TESTS
X = np.array([[ 0.27170746,  0.89441607,  0.64849028],
              [ 0.42296173,  0.54342876,  0.47889235],
              [ 0.48688657,  0.11082849,  0.10691689],
              [ 0.04419385,  0.68777309,  0.49437059],
              [ 0.70143641,  0.09964604,  0.20949214],
              [ 0.01725016,  0.37424641,  0.94070338]])

Y = np.array([[ 0.24232741,  0.08413896,  0.014919  ],
              [ 0.15801316,  0.31713579,  0.0416702 ],
              [ 0.15784176,  0.50998073,  0.45405793],
              [ 0.44382259,  0.44515729,  0.49186482],
              [ 0.00695024,  0.23603969,  0.77601819]])

V = nearest_redi(X,Y)
V_ans = np.array([2, 3, 0, 2, 0, 4])

assert np.array_equal(V, V_ans)
### END HIDDEN TESTS
```

### Part 2: Find the number of Redis inside a radius of $r$ from each Homo BITSian

<!-- grade_id: num_redis, points: 0, grade: false, solution: true -->
```python
def num_redis(X,Y,r):
    """
    Given matrices X and Y, the function returns a vector, where the i-th element 
    of the vector contains the number of Redis inside a radius of r,
    centered at the i-th Homo BITSian.
    
    Parameters: X,Y,r
    Returns: W
    """
    
    D = distances(X,Y)
    
    ### BEGIN SOLUTION
    bool_less_r = D < r
    W = np.sum(bool_less_r, axis = 1)
    ### END SOLUTION
    
    return W
```

<!-- grade_id: visible-task2, points: 0, grade: true, solution: false -->
```python
print("Running base test case 1...")

W_test1 = num_redis(X_test1, Y_test1, 5.)
W_ans_test1 = np.array([1, 0, 2])

assert np.array_equal(W_test1, W_ans_test1)

print("Base test case 1 successful!!\n")



print("Running base test case 2...")

W_test2 = num_redis(X_test2, Y_test2, 8.)
W_ans_test2 = np.array([1, 2])

assert np.array_equal(W_test2, W_ans_test2)

print("Base test case 2 successful!!\n")
```

Output:
```
Running base test case 1...
Base test case 1 successful!!

Running base test case 2...
Base test case 2 successful!!
```

<!-- grade_id: hidden_task2, points: 3, grade: true, solution: false -->
```python
# Running hidden test case for Part 2. Don't edit the cell.                                     *** 3 marks ***
### BEGIN HIDDEN TESTS

W = num_redis(X,Y, 0.7)
W_ans = np.array([2, 5, 4, 4, 3, 3])

assert np.array_equal(W, W_ans)

### END HIDDEN TESTS
```